# <span style="color:teal;">CIS 211 Project 6:  N-Body GUI</span>

##### Due 11:00 P.M. Tuesday Feb 28, 2017

##### Reading:  Online Tutorials

For this week's project you will create the elements used in a GUI for visualizing the solar system simulation: 
* a subclass of Body that adds attributes used to display a planet 
* a new type of Canvas with methods for drawing and moving planets
* a frame that contains a label and a spinbox used to control the number of planets to display
* another frame that has a run button and a place to enter the simulation parameters

###  <span style="color:teal">N-Body Simulation Code</span>

This week's project will use the Vector and Body classes from Project 4 and the `step_system` function from Project 5.  There are two ways to use that code in this project:  (1) copy your definitions from the previous projects, or (2) download the instructor's solution from Canvas.

Instructions for both methods are in the next two sections.  **Use either of these methods, but not both.**  

Note: you will earn **extra credit points** if you use your own code.

####  <span style="color:teal">Option 1: &nbsp; Use Your Own Code</span>

Choose this method only if your Vector and Body class passed all their unit tests and the `step_system` function passed its tests. Copy the complete definitions from your other notebooks and paste them into the code cells below.  

```python
# Copy and paste a working version of the Vector class into this cell.

pass
```

```python
# Copy and paste a working version of the Body class into this cell.

pass
```

```python
# Copy and paste a working version of the step_system function into this cell.

pass
```

####  <span style="color:teal">Option 2: &nbsp; Download `Orbits.pyc` from Canvas </span>

Download one of the `Orbits.cpython-3x.pyc` files, rename it `Orbits.pyc`, and move it to the same folder as this notebook.  Then uncomment and execute the `import` command in the code cell below.

**Make sure you get the right file for your version of Python.**

```python
from Orbits import *
```

###  <span style="color:teal">Libraries</span>

Execute these code cells each time you open the notebook.  The cell that defines two global constants will allow us to run autograder tests when the notebook is loaded into Jupyter but not when the program is run from the command line.

```python
IPython = (__doc__ is not None) and ('IPython' in __doc__)
Main    = __name__ == '__main__'
```

```python
if IPython:
    %gui tk
```

```python
import tkinter as tk
import tkinter.ttk as ttk
from tkinter import filedialog
from time import sleep
```

###  <span style="color:teal">Helper Function</span>

To help test and debug your code execute this code cell to define a function named `read_bodies`.

The first argument passed to the function should be the name of the file containing planet descriptions.  Download `solarsystem.txt` from Canvas and save it in the same directory as this notebook.

The second argument is the type of Body object you want to make.  If you call
```
>>> read_bodies('solarsystem.txt', Body)
```
you will get back a list of Body objects (you should to this to make sure you can read the data file).  

Part 1 of this week's project is to defined a new class named TkBody that is derived from the Body class. If you call
```
>>> read_bodies('solarsystem.txt', TkBody)
```
you will get a list of TkBody objects that include the attributes needed to draw the planets on a canvas.

```python
def read_bodies(filename, cls):
    '''
    Read descriptions of planets, return a list of body objects.  The first
    argument is the name of a file with one planet description per line, the
    second argument is the type of object to make for each planet.
    '''
    if not issubclass(cls, Body):
        raise TypeError('cls must be Body or a subclass of Body')

    bodies = [ ]

    with open(filename) as bodyfile:
        for line in bodyfile:
            line = line.strip()
            if len(line) == 0 or line[0] == '#':
                continue
            name, m, rx, ry, rz, vx, vy, vz, diam, color = line.split()
            args = {
                'mass' : float(m),
                'position' : Vector(float(rx), float(ry), float(rz)),
                'velocity' : Vector(float(vx), float(vy), float(vz)),
            }
            opts = {'name': name, 'color': color, 'size': int(diam)}
            for x in opts:
                if getattr(cls, x, None):
                    args[x] = opts[x]
            bodies.append(cls(**args))

    return bodies

```

###  <span style="color:teal">Part 1: &nbsp; TkBody (20 points)</span>

Fill in the definition of the TkBody class in the code cell below.  The constructor should accept the same arguments as the Body class constructor plus two new attributes:
* `size` is the radius, in pixels, of the circle to draw to represent the body
* `color` is a string used to defined the fill color of the circle

Define two getter methods, named `size` and `color`, to return the values of those attributes.

A third new attribute of a TkBody object is named `graphic`.  It will hold the ID of the circle created when the body is drawn. Initialize this attribute to None when the body is created, and define a "setter" named `set_graphic` and a "getter" named `graphic`.

**Note:** If you need a reminder of the arguments passed to Body type `?Body` or `help(Body)` in a Jupyter code cell.


##### <span style="color:red">Code:</span>

<!-- grade_id: tkbody, points: 10, grade: true, solution: true -->
```python
class TkBody(Body):
    def __init__(self, mass, position, velocity, name ,size, color, graphic=None):
        Body.__init__(self, mass, position, velocity, name)
        self._graphic = graphic
        self._size = size
        self._color = color
    def size(self):
        return self._size
    def color(self):
        return self._color
    def graphic(self):
        return self._graphic
    def set_graphic(self, new):
        self._graphic = new
```

##### <span style="color:red">Tests:</span>

Use the following code cell as a "sandbox" if you want to do your own tests of the TkBody class.  You can add additional cells here if you want.

```python
c = TkBody(0, Vector(0,0,0), Vector(0,0,0), '', 0, 'black')
print(c.color())
print(c.size())
print(c.name())
print(bodies[0].name)
```

Output:
```
black
0

TypeError                                 Traceback (most recent call last)
<ipython-input-69-f17fa78e3d76> in <module>()
      3 print(c.size())
      4 print(c.name())
----> 5 print(bodies[0].name)

TypeError: 'NoneType' object is not subscriptable
```

##### <span style="color:red">Autograder Tests:</span>

**Important:** &nbsp;  the code cells in this section will be used by `nbgrader` to run automated tests.  Do not move, delete or alter these cells in any way.

<!-- grade_id: tkbody_test1, points: 2, grade: true, solution: false -->
```python
if IPython:
    b = TkBody(0, Vector(0,0,0), Vector(0,0,0), '', 0, 'black')

    assert isinstance(b, Body)
    assert isinstance(b, TkBody)

    assert b.size() == 0
    assert b.color() == 'black'
    assert b.graphic() is None

    b.set_graphic(0)
    assert b.graphic() == 0
```

<!-- grade_id: tkbody_test2, points: 3, grade: true, solution: false -->
```python
if IPython:
    bodies = read_bodies('solarsystem.txt', TkBody)

    assert isinstance(bodies[0], TkBody)

    assert bodies[0].name() == 'sun'
    assert bodies[0].size() == 10
    assert bodies[0].color() == '#ffff00'
```

##### <span style="color:red">Documentation:</span>

<!-- grade_id: tkbody_doc, points: 5, grade: true, solution: true -->
I started doing this with name=None at the end of the constructor of tkbody, because it was optional in Body i thought it would have to be optional in tkbody, but then in order to get 'graphic' to be in the constructor it had to optional as well so i set self._graphic=None._ so, after much tinkering i just switched name and graphic around and it worked somehow so im happy with it

###  <span style="color:teal">Part 2: &nbsp; SolarSystemCanvas (20 points)</span>

A SolarSystemCanvas is a type of Canvas used to draw planets.  An instance of this class has the following methods:
* `set_planets` is passed a list of all the bodies in the simulation; it should save this list in an instance variable for use later
* `view_planets` is passed an integer `n`, and it should draw circles representing the first `n` bodies in the list (and save the value of `n` in an instance variable)
* `move_planets` is also passed a complete list of body objects; it should use the positions of the first `n` bodies to compute new locations of the circles of the `n` planets currently on the canvas

Much of the code for this class has been written already for you.  Here are the details of the methods you need to fill in.

#### `_compute_scale` 

An important role for this class is to convert solar system coordinates (billions of meters) into screen coordinates (hundreds of pixels).  The method named `_compute_scale` calculates a scaling factor so that when a body is displayed it is placed at the correct location on the screen.  The formula is

$$f = \frac{p\,/\,2}{| \, d_\mathrm{max} \, |} $$

where $p$ is the height or width (whichever is smaller) of the canvas, in pixels, and $d_\mathrm{max}$ is the largest distance to any planet in the list passed to the method.

For example, suppose the screen is 500 pixels wide and 400 pixels tall, and this method is called with a list of 6 bodies, for the sun and all the planets up through Jupiter.  The $x$ coordinate of Jupiter at the start of the simulation is $-7.5 \times 10^{11}$, which defines a scale factor of $2.66 \times 10^{-10}$.

There are different ways to figure out the largest distance.  A simple method is to find the maximum of the norm of each position vector; another is to find the largest absolute value of any $x$ or $y$ coordinate.  

You can also include a "fudge factor" if you want to add a little extra space so the furthest orbit does not bump up against the edge of the canvas.  For example, if you use $1.1 \times d_\mathrm{max}$ in the denominator the planets will be about 10% closer to the center.

Your method should compute the scale factor and save it in an instance variable named `_scale` so it can be used later when the planets are drawn and moved.

**Important:** You need to save the scale factor in `self._scale` because that is the name used by the other methods which we have written for you.

#### `view_planets` 

This method will be passed an integer `n`, and it needs to draw circles representing the first `n` planets in the current list of TkBody objects (which have been saved in the instance variable `self._planets`).  Your method needs to:
* compute the scale factor for this list of planets
* delete any planets and lines that are on the canvas
* draw circles for the first `n` planets in `self._planets`
* save the value of `n` in `self._outer` so it can be used by `move_planets`
The circles should use the size and color attributes of the TkBody objects that represent the planets.

#### `move_planets` 

This method will be called by the main simulator after each time step.  The argument passed to the method is the list of TkBody objects with their updated position vectors.  For each planet currently on the canvas:
* get the $x$ and $y$ coordinates of the current location of the circle
* compute the $x$ and $y$ coordinates for the updated location
* move the circle to the new location
* draw a line segment from the old location to the new location

Two methods that compute screen locations have been written for you.  Call `self._current_loc(p)` to get the $x$ and $y$ coordinates of the center of the circle representing planet `p`.  Call `self._compute_loc(p)` to get the $x$ and $y$ coordinates for the new center of the circle based on the position vector of planet `p`.

<!-- grade_id: solar_system_canvas, points: 10, grade: true, solution: true -->
```python
class SolarSystemCanvas(tk.Canvas):
    
    def __init__(self, parent, height=600, width=600):
        tk.Canvas.__init__(self, parent, height=height, width=width, background='gray90', highlightthickness=0)
        self._planets = None
        self._outer = None
        self._scale = None
        self._offset = Vector(int(self['width'])/2, int(self['height'])/2, 0)
        self._height = height
        self._width = width
        
    def set_planets(self, lst):
        self._planets = lst
        self._outer = len(lst)
        self._compute_scale(lst)
        self.view_planets(len(lst))
        
    def view_planets(self, n):
        self.delete('all')
        for planet in self._planets[:n]:
            radius = planet.size()
            x, y = self._compute_loc(planet)
            gp = self.create_oval((x-radius , y-radius),(x+radius , y+radius), fill=planet.color())
            planet.set_graphic(gp)
        self._outer = n
        
    def move_planets(self, lst):
        for planet in lst[:self._outer]:
            initial_center = self._compute_loc(planet)
            planet.move(1)
            post_center = self._compute_loc(planet)
            center_delta_x = post_center[0] - initial_center[0]
            center_delta_y = post_center[1] - initial_center[1]
            self.move(planet.graphic(), center_delta_x, center_delta_y)
            self.create_line((initial_center),(post_center))
            sleep(0.025)
            self.update_idletasks()
            
    def _compute_scale(self, bodies):
        p = min((self._height, self._width))
        dm = 0
        for b in bodies:
            if b.position().norm() > dm:
                dm = b.position().norm()
        self._scale = ((p/2)/(1.1 * dm))
    
    def _compute_loc(self, p):
        pos = p.position() * self._scale + self._offset
        return pos.x(), pos.y()
    
    def _current_loc(self, p):
        ul, ur, _, _ = self.coords(p.graphic())
        return ul + p.size(), ur + p.size()
 
```

##### <span style="color:red">Tests:</span>

Use the following code cell as a "sandbox" if you want to do your own tests of the ViewFrame class.  You can add additional cells here if you want.

```python

```

##### <span style="color:red">Autograder Tests:</span>

**Important:** &nbsp;  the code cells in this section will be used by `nbgrader` to run automated tests.  Do not move, delete or alter these cells in any way.

```python
if IPython:
    canvas = SolarSystemCanvas(tk.Tk(), height=400, width=400)
    canvas.pack()
    bodies = read_bodies('solarsystem.txt', TkBody)
```

<!-- grade_id: canvas_test1, points: 2, grade: true, solution: false -->
```python
if IPython:
    canvas.set_planets(bodies)
    canvas.view_planets(2)
    assert [canvas.type(x) for x in canvas.find_all()].count('oval') == 2
```

<!-- grade_id: canvas_test2, points: 3, grade: true, solution: false -->
```python
if IPython:
    for i in range(10):
        step_system(bodies)
        canvas.move_planets(bodies)
    tk_objects = canvas.find_all()
    assert len(tk_objects) == 22
    assert [canvas.type(x) for x in tk_objects].count('line') == 20
```

##### <span style="color:red">Documentation:</span>

<!-- grade_id: canvas_doc, points: 5, grade: true, solution: true -->
This one had a lot of context to it that i struggled with, compute scale went along pretty easily by following instructions view planets was not too hard, just had to to give graphic a value each time (took me a while to figure out), move planets was an emotional rodeo but i think it turned out pretty neat, it took me a while to figure out how to move the circles and make the lines, i messed around with the time steps a little to arrive where it is now

###  <span style="color:teal">Part 3: &nbsp; ViewControl (20 points)</span>

A ViewControl component allows the user to specify the number of planets to display on the canvas.  It's basically just a spinbox with a label, but with the extra behavior that the spinbox value can only between 2 (the canvas should always show at least the Sun and Mercury) and however many bodies are in the data set.

The constructor will be passed a reference to a callback function.  You should pass this function to the constructor that makes the Spinbox widget inside the ViewFrame.

You need to fill in the definition of the ViewFame class:
* the constructor should create a Label and a Spinbox and add them to this frame
* the `state` attribute of the Spinbox should initially be `tk.DISABLED` since we don't know yet what the maximum value is
* the `reset` method will be passed an integer `n`; it should set the spinbox lower limit to 2, set the upper limit to `n`, set the current value to `n`, and enable the spinbox by setting its state to `tk.NORMAL`
* the `nbodies` method should return 0 if the Spinbox is disable, otherwise it should return the current value in the spinbox, converted to an integer.

**Important:** &nbsp; In order to pass autograder tests for ViewControl make sure you save a reference to the Spinbox in an instance variable named `_spinbox`

##### <span style="color:red">Code:</span>

<!-- grade_id: view_control, points: 10, grade: true, solution: true -->
```python
class ViewControl(tk.Frame):
    
    def __init__(self, parent, callback):
        tk.Frame.__init__(self, parent)
        
        self['width'] = 100             # <- the tests do not rely on these
        self['height'] = 25             # <- assignments so you can delete
        self['background'] = 'gray'     # <- them if you wish
        self._callback = callback
        
        Label = tk.Label(self, text='View:')
        Label.grid(row=0, column=0)
        
        self._spinbox = tk.Spinbox(self, from_=2, to=8, state=tk.DISABLED, command=self._callback)
        self._spinbox.grid(row=0, column =1)
        
        n = 0
    def reset(self, nbodies):
        self._spinbox.config(to=nbodies, state=tk.NORMAL)
        self._spinbox.delete(0,"end")
        self._spinbox.insert(0,nbodies)
        n = 1
    def nbodies(self):
        if n != 0:
            return 0
        else:
            return(_spinbox.get())

```

##### <span style="color:red">Tests:</span>

Use the following code cell as a "sandbox" if you want to do your own tests of the ViewFrame class.  You can add additional cells here if you want.

```python

```

##### <span style="color:red">Autograder Tests:</span>

**Important:** &nbsp;  the code cells in this section will be used by `nbgrader` to run automated tests.  Do not move, delete or alter these cells in any way.

```python
def vf_cb():
    global called
    called = True

if IPython:
    vf = ViewControl(tk.Tk(), vf_cb)
    vf.pack()        
```

<!-- grade_id: vf_test_1, points: 2, grade: true, solution: false -->
```python
if IPython:
    counts = { }
    for x in vf.children.values():
        counts.setdefault(type(x).__name__, 0)
        counts[type(x).__name__] += 1
    assert counts == { 'Label': 1, 'Spinbox': 1 }
```

<!-- grade_id: vf_test_2, points: 3, grade: true, solution: false -->
```python
if IPython:
    called = False
    vf._spinbox.invoke('buttondown')
    vf._spinbox.invoke('buttonup')
    assert not called
    
    vf.reset(5)
    assert vf._spinbox.get() == '5'
    assert vf._spinbox['from'] == 2
    assert vf._spinbox['to'] == 5

    vf._spinbox.invoke('buttondown')
    assert vf._spinbox.get() == '4'
    vf._spinbox.invoke('buttonup')
    assert vf._spinbox.get() == '5'
    
    assert called
```

##### <span style="color:red">Documentation:</span>

<!-- grade_id: view_doc, points: 5, grade: true, solution: true -->
Well i created a label and a spinbox, the spinbox is by default disabled and it follows the label. for setting the spinbox i used config, as well as enabling it. n is just a sentinel to check if the spinbox is enabled. The delete and insert business is just to make it so the spinbox starts with the nbodies value.

###  <span style="color:teal">Part 4: &nbsp; RunFrame (20 points)</span>

This project is similar to the first one, but here you will fill out the definition of a class called RunFrame.  A RunFrame is a frame that has a run button, text entry boxes for the time step size and number of steps to run, and a progress bar to show how many steps have been executed.  

The entry boxes for the time step size and number of steps should be initialized to show the default values for those parameters:  86459 for the time step size and 365 for the number of steps.

After an application has created a RunFrame object named `r` it should be able to call the following methods to interact with `r`:
* `r.dt()` returns the time step size from the time step entry box
* `r.nsteps()` returns the number of time steps indicated in that entry box
* `t.target(n)` sets the upper bound in the progress bar
* `t.update_progress(n)` sets the value to show in the progress bar

The constructor will be passed a reference to a callback function that will be called when the run button is pushed.  You should pass this function to the constructor that makes the Button.

**Note:** &nbsp; In order to test your RunFrame we need to know the names you give to some of the parts of the frame.  Make sure you use the following instance variable names:
* `_run_button` for the button that starts the simulation
* `_dt_entry` for the entry box for time step size
* `_nsteps_entry` for the entry box with the number of steps

##### <span style="color:red">Code:</span>

**Important:** &nbsp; Add your Python code to the following code cell.  Do not delete, rename, or copy this cell.

<!-- grade_id: run_frame, points: 5, grade: true, solution: true -->
```python
class RunFrame(tk.Frame):
    
    def __init__(self, parent, callback):
        tk.Frame.__init__(self, parent)
        
        self['width'] = 200
        self['height'] = 100
        self['background'] = 'gray'
        
        self._run_button = tk.Button(self, text='Run', command=callback)
        self._run_button.grid(row=0, column=0)
        
        self._dt_label = tk.Label(self, text='dt')
        self._dt_label.grid(row=0, column=1)
        
        self._dt_entry = tk.Entry(self)
        self._dt_entry.grid(row=0, column=2, columnspan=2)
        self._dt_entry.insert(0,'86459')

        
        self._nsteps_label = tk.Label(self, text='nsteps')
        self._nsteps_label.grid(row=1, column=1)
        
        self._nsteps_entry = tk.Entry(self)
        self._nsteps_entry.grid(row=1, column=2, columnspan=2)
        self._nsteps_entry.insert(0,'365')
        
        self._progress = ttk.Progressbar(self, orient='horizontal', length=200, mode='determinate')
        self._progress.grid(row=3, column=0, columnspan=3)

        
    def dt(self):
        return int(self._dt_entry.get())

    def nsteps(self):
        return int(self._nsteps_entry.get())
        
    def init_progress(self, n):
        self._progress.config(maximum=n)
        
    def update_progress(self, n):
        m = self._progress['value']+n
        self._progress.config(value=m)
        
    def clear_progress(self):
        self._progress['value'] = 0

```

##### <span style="color:red">Tests:</span>

Use the following code cell as a "sandbox" if you want to do your own tests.  You can add additional cells here if you want.

```python

```

##### <span style="color:red">Autograder Tests:</span>

**Important:** &nbsp;  the code cells in this section will be used by `nbgrader` to run automated tests.  Do not move, delete or alter these cells in any way.

```python
def rf_cb():
    global calls
    calls += 1

if IPython:
    rf = RunFrame(tk.Tk(), rf_cb)
    rf.pack()
```

<!-- grade_id: rf_test1, points: 3, grade: true, solution: false -->
```python
if IPython:
    counts = { }
    for x in rf.children.values():
        counts.setdefault(type(x).__name__, 0)
        counts[type(x).__name__] += 1
    assert counts['Button'] == 1
    assert counts['Entry'] == 2
    assert counts['Progressbar'] == 1
    assert rf._dt_entry.get() == '86459'
    assert rf._nsteps_entry.get() == '365'
```

<!-- grade_id: rf_test2, points: 2, grade: true, solution: false -->
```python
if IPython:
    calls = 0
    rf._run_button.invoke()
    rf._run_button.invoke()
    assert calls == 2
```

<!-- grade_id: rf_test3, points: 3, grade: true, solution: false -->
```python
if IPython:
    rf._dt_entry.delete(0, tk.END)
    rf._dt_entry.insert(0, '100')
    assert rf.dt() == 100
    rf._nsteps_entry.delete(0, tk.END)
    rf._nsteps_entry.insert(0, '1000')
    assert rf.nsteps() == 1000
```

<!-- grade_id: rf_test4, points: 2, grade: true, solution: false -->
```python
if IPython:
    rf.init_progress(100)
    for i in range(5):
        rf.update_progress(10)
    assert rf._progress['value'] == 50
```

##### <span style="color:red">Documentation:</span>

<!-- grade_id: run_doc, points: 5, grade: true, solution: true -->
this one was the easiest of them all, we made the grid and set the entries default values. the rogressbar was made and accumulated the values accordingly and dt and nsteps return integers

###  <span style="color:teal">Top Level Application</span>

The application that instantiates all the widgets and links them together is shown below. You do not need to write any code -- the program will work when all of your widget definitions are complete. If you execute this code cell in the Jupyter notebook you will get a working application that will allow you to load a set of planet definitions and run the simulatio.

Even though you do not have to write anything you should look closely at this code since it will help you understand the methods you define in each of your components.

To test your GUI, we will export the notebook to a command line application and run the application.  The graders will run the program and check the state of the canvas at the end of the simulation.

#### `load_cb` 

This function is called when the user clicks the Load button.  It uses a dialog to get the name of the file with planet definitions and reads the contents of the file to get a list of TkBody objects.  Notice how:
* the list is passed to the `set_planets` method in the SolarSystemCanvas widget, which will draw circles for all the planets
* the number of bodies is passed to the `reset` method in the ViewControl, which allows the user to click the up/down buttons to change the number of planets displayed

#### `view_cb` 

This function is the callback for the ViewControl component.  Whenever the user changes the value of the counter the new value is passed to `view_planets` so the canvas is updated to show the new number of planets.

#### `run_cb` 

This function is the one that runs the simulation.  It is invoked when the user clicks the Run button in the RunFrame widget.  It gets the simulation parameters from the widget and then calls the `time_step` function to start the simulation.

There body of the `time_step` function could just have a for loop that calls `step_system` for the specified number of time steps.  But since `run_cb` is a callback, Tk won't update the display until the callback is done, which means we won't see the planets move until after the last time step.

The technique implemented here is pretty common in GUI programs.  It runs one time step and then uses the `after_idle` function to schedule a call to run the next time step 0.02 seconds later.  Once the next call is scheduled the callback exits and returns control to the GUI and we will see the planets moving.

```python
root = tk.Tk()
root.title("Solar System")

bodies = None

def load_cb():
    global bodies
    fn = tk.filedialog.askopenfilename()
    bodies = read_bodies(fn, TkBody)
    canvas.set_planets(bodies)
    view_count.reset(len(bodies))

def view_cb():
    canvas.view_planets(view_count.nbodies())
    
def run_cb():
    
    def time_step():
        nonlocal nsteps
        step_system(bodies, dt)
        canvas.move_planets(bodies)
        run_frame.update_progress(1)
        sleep(0.02)
        if nsteps > 0:
            nsteps -= 1
            canvas.after_idle(time_step)
        else:
            run_frame.clear_progress()
        
    nsteps = run_frame.nsteps()
    run_frame.init_progress(nsteps)
    dt = run_frame.dt()
    canvas.after_idle(time_step)

canvas = SolarSystemCanvas(root)
canvas.grid(row = 0, column = 0, columnspan = 3, padx = 10, pady = 10, sticky="nsew")

tk.Button(root, text='Load', command=load_cb).grid(row=1, column=0, pady = 20)

view_count = ViewControl(root, view_cb)
view_count.grid(row=1, column=1, pady=20)

run_frame = RunFrame(root, run_cb)
run_frame.grid(row=1, column=2, pady=20)

if Main and not IPython:
    try:
        bodies = read_bodies("solarsystem.txt", TkBody)
        canvas.set_planets(bodies)
        view_count.reset(len(bodies))
        for i in range(5):
            view_count._spinbox.invoke('buttondown')
        run_frame._nsteps_entry.delete(0, tk.END)
        run_frame._nsteps_entry.insert(0,'100')
        root.update()
        run_frame._run_button.invoke()
    except Exception as err:
        print(err)
    input('hit return to continue...')

```

Output:
```
NameError                                 Traceback (most recent call last)
<ipython-input-62-b161955996aa> in <module>()
     33     canvas.after_idle(time_step)
     34 
---> 35 canvas = SolarSystemCanvas(root)
     36 canvas.grid(row = 0, column = 0, columnspan = 3, padx = 10, pady = 10, sticky="nsew")
     37 

NameError: name 'SolarSystemCanvas' is not defined
```

##### <span style="color:red">Simulation Score</span>

**Important:** &nbsp; You do not have to write anything for this part of the project.  Leave the cell below empty, in spite of the fact that is says "your documentation here".

The graders will use the cell to write comments about your GUI and to enter the grade for your solar system simulation.  **Do not edit, move, rename, or delete the following cell.**

<!-- grade_id: simulation, points: 20, grade: true, solution: true -->
YOUR DOCUMENTATION HERE